// intro2.rs
//
// Make the code print a greeting to the world.
//
// Execute `rustlings hint intro2` or use the `hint` watch subcommand for a
// hint.

// I AM NOT DONE

fn main() {
    println!("Hello there!");
    println!("{number:>5}", number=1);
    println!("{number:0>width$}", number=1, width=5);
}
